# Display Exercise 1

## Imports

Put any needed imports needed to display rich output the following cell:

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
from IPython.display import display
from IPython.display import Image
from IPython.display import HTML
```

<!-- grade_id: displayex01a, points: 2, grade: true, solution: false -->
```python
assert True # leave this to grade the import statements
```

# Basic rich display

Find a Physics related image on the internet and display it in this notebook using the `Image` object.

* Load it using the `url` argument to `Image` (don't upload the image to this server).
* Make sure the set the `embed` flag so the image is embedded in the notebook data.
* Set the width and height to `600px`.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
Image(url='https://english.tau.ac.il/sites/default/files/styles/reaserch_main_image_580_x_330/public/sackler%20physics%20cropped.jpg?itok=oanzfnK-')
```

Output:
```
<IPython.core.display.Image object>
```

<!-- grade_id: displayex01b, points: 4, grade: true, solution: false -->
```python
assert True # leave this to grade the image display
```

Use the `HTML` object to display HTML in the notebook that reproduces the table of Quarks on [this page](http://en.wikipedia.org/wiki/List_of_particles). This will require you to learn about how to create HTML tables and then pass that to the HTML object for display. Don't worry about styling and formatting the table, but you should use LaTeX where appropriate.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
s = """<table>
<tr>
<th>Header 1</th>
<th>Header 2</th>
</tr>
<tr>
<td>row 1, cell 1</td>
<td>row 1, cell 2</td>
</tr>
<tr>
<td>row 2, cell 1</td>
<td>row 2, cell 2</td>
</tr>
</table>"""
```

```python
h = HTML(s)
display(h)
```

Output:
```
<IPython.core.display.HTML object>
```

<!-- grade_id: displayex01c, points: 4, grade: true, solution: false -->
```python
assert True # leave this here to grade the quark table
```